```python
# Initialize Otter
import otter
grader = otter.Notebook("hw2.ipynb")
```

## Notes:
- Please don't import external packages I haven't imported as part of the assignment, as this will cause otter grader to crash upon submission, resulting in a zero, this is fixable but not ideal.
- To turn in the assignment, run all code and test, save the notebook, then run the final cell to generate a .zip file to submit on gradescope.
- AI tools are allowed, but please disclose the way you used AI tools in the AI disclosure section (ungraded). 

# Homework 2

```python
# Setup
# You may import math if needed in your solutions.
import math
```

## **Question 1**: Conditional Logic - Projectile Motion Analysis

### Q1a: Projectile Launch Classification
A projectile is launched with initial velocity components v_x (m/s) and v_y (m/s). Write a function `classify_projectile(v_x, v_y)` that returns a string classification:
- "vertical_up" if v_x = 0 and v_y > 0
- "vertical_down" if v_x = 0 and v_y < 0  
- "horizontal" if v_x ≠ 0 and v_y = 0
- "oblique_up" if v_x ≠ 0 and v_y > 0
- "oblique_down" if v_x ≠ 0 and v_y < 0
- "stationary" if v_x = 0 and v_y = 0

```python
def classify_projectile(v_x, v_y):
    # Write your code here!
```

```python
grader.check("q1a")
```

### Q1b: Range and Flight Time Calculator
Building on Q1a, write a function `calculate_projectile_motion(v_x, v_y, g=9.8)` that calculates projectile motion parameters. Use your `classify_projectile` function from Q1a and return different values based on the projectile type:

- For "stationary": return (0, 0) for (range, flight_time)
- For "vertical_up": return (0, flight_time) where flight_time = 2*|v_y|/g
- For "vertical_down": return (0, 0) since it hits the ground immediately at launch height
- For "horizontal": return (0, 0) since it hits the ground immediately at launch height
- For "oblique_up": return (range, flight_time) where:
  - flight_time = 2*v_y/g  
  - range = v_x * flight_time
- For "oblique_down": return (0, 0) since it hits the ground immediately

Return a tuple (range, flight_time).

```python
def calculate_projectile_motion(v_x, v_y, g=9.8):
    # Write your code here!
    # Hint: Use your classify_projectile function from Q1a
```

```python
grader.check("q1b")
```

### Q1c: Safety Assessment
Building on your previous functions, write `assess_projectile_safety(v_x, v_y, max_range=100, max_height=50, g=9.8)` that determines if a projectile launch is safe. Use your previous functions and return a string:

- "SAFE" if:
  - The projectile is "stationary", OR
  - For "vertical_up": maximum height ≤ max_height (max_height = v_y²/(2g)), OR
  - For "vertical_down": always return "SAFE" (hits ground immediately), OR  
  - For "horizontal": always return "SAFE" (hits ground immediately), OR
  - For "oblique_up": range ≤ max_range AND max_height ≤ max_height, OR
  - For "oblique_down": always return "SAFE" (hits ground immediately)

- "UNSAFE" otherwise

- For any trajectory that exceeds limits, return "UNSAFE"

```python
def assess_projectile_safety(v_x, v_y, max_range=100, max_height=50, g=9.8):
    # Write your code here!
    # Hint: Use your classify_projectile and calculate_projectile_motion functions
```

```python
grader.check("q1c")
```

## **Question 2**: For Loops - Oscillations and Wave Analysis

### Q2a: Simple Harmonic Motion Analysis
Write a function `analyze_oscillation(amplitude, num_points, periods=1, T=1.0)` that calculates the position of a simple harmonic oscillator at equally spaced time intervals over the given number of `periods`.

The position is given by: x(t) = amplitude * cos(2π * t / T) where T is the period.

The time should range in even spaces between 0 and (periods * T) - (periods * T)/num_points

Return a list of position values (floats). (`.append()` within the for loop can be useful here)

```python
def analyze_oscillation(amplitude, num_points, periods=1, T=1.0):
    # Write your code here!
    # Use a for loop with range(num_points)
```

```python
grader.check("q2a")
```

### Q2b: Wave Interference Analysis
Write a function `analyze_wave_interference(frequencies, amplitudes, time=0)` that calculates the total displacement from multiple interfering waves at a given time.

Each wave follows: y_i(t) = A_i * sin(2π * f_i * t) where A_i is amplitude and f_i is frequency.

The function takes:
- `frequencies`: a list of wave frequencies (Hz)
- `amplitudes`: a list of corresponding amplitudes (same length as frequencies)  
- `time`: time at which to evaluate (default 0)

Use a for loop to iterate over the frequency and amplitude lists simultaneously and sum all wave contributions.

Return the total displacement (float).

If empty lists are passed, the function should return 0 because there is no wave displacement.

```python
def analyze_wave_interference(frequencies, amplitudes, time=0):
    # Write your code here!
```

```python
grader.check("q2b")
```

If you accidentally make an infinite loop or infinite recursion in the following questions, you may be unable to run more cells. If this happens, restart the kernel, and if that doesn't fix it, save the notebook, and click the assignment link again on bCourses to start a new datahub session. 

## **Question 3**: Recursive Functions

### Example: Recursive Factorial Function

Recursion is a programming technique where a function calls itself to solve a problem by breaking it down into smaller, similar subproblems. A recursive function typically has two key components:

1. **Base case**: A condition that stops the recursion
2. **Recursive case**: The function calls itself with a modified input

Here's an example of a recursive function to compute factorials:

```python
def factorial(n):
    # Base case: factorial of 0 or 1 is 1
    if n <= 1:
        return 1
    # Recursive case: n! = n * (n-1)!
    else:
        return n * factorial(n - 1)
```

This function works by recognizing that `n! = n × (n-1)!`, and continues calling itself with smaller values until it reaches the base case.

### Q3a: Recursive Fibonacci Function

Write a recursive function `fibonacci(n)` that returns the nth Fibonacci number. The Fibonacci sequence is defined as:
- F(0) = 0
- F(1) = 1  
- F(n) = F(n-1) + F(n-2) for n ≥ 2

Your function should handle the base cases (n=0 and n=1) and use recursion for larger values.

Return an integer.

```python
def fibonacci(n):
    # Write your code here!
```

```python
grader.check("q3a")
```

### Q3b: Recursive List Sum

Write a recursive function `recursive_sum(numbers)` that calculates the sum of all elements in a list of numbers using recursion.

Your function should:
- **Base case**: If the list is empty, return 0
- **Recursive case**: Return the first element plus the sum of the rest of the list

Hint: You can use `numbers[1:]` to get the list without the first element.

Return a number (int or float).

```python
def recursive_sum(numbers):
    # Write your code here!
```

```python
grader.check("q3b")
```

## Question 4: While Loops - Physical System Simulation

In physics, while loops are commonly used to simulate systems that evolve over time until certain conditions are met. This question explores two important applications of while loops in physics computations.

---

### Q4a: Radioactive Decay Simulation

Write a function `radioactive_decay(initial_atoms, half_life, threshold)` that simulates radioactive decay using a while loop. The function should:

- Start with `initial_atoms` radioactive atoms
- Use a while loop that continues **while the number of atoms is greater than the threshold**
- In each time step (representing one half-life period), reduce the number of atoms by half
- Count how many half-life periods it takes to reach the threshold
- Return a tuple `(final_atoms, time_periods)` where `final_atoms` is the remaining atoms and `time_periods` is the number of half-life periods elapsed

**Physics Note**: Radioactive decay follows the equation N(t) = N₀ × (1/2)^(t/t₁/₂), where N₀ is initial atoms, t is time, and t₁/₂ is the half-life.

Return a tuple of (int, int).

```python
def radioactive_decay(initial_atoms, half_life, threshold):
    # Write your code here!
```

```python
grader.check("q4a")
```

### Q4b: Chaotic System Analyzer

In physics, chaotic systems exhibit sensitive dependence on initial conditions. The logistic map is a famous example: `x_next = r * x * (1 - x)`, where different values of `r` lead to different behaviors.

First, we provide you with a function that implements the logistic map:

```python
def logistic_map(x, r):
    """Apply one iteration of the logistic map: x_next = r * x * (1 - x)"""
    return r * x * (1 - x)
```

Now write a function `find_fixed_point(initial_x, r, tolerance=0.001, max_iterations=1000)` that uses the logistic map to find when the system reaches a fixed point (or gets close to one). The function should:

- Start with the given `initial_x` value
- Use `while True:` to create an infinite loop
- In each iteration, apply the logistic map: `x = logistic_map(x, r)`
- **Break the loop** when either:
  - The change between consecutive iterations is less than `tolerance` (system has stabilized)
  - OR when `max_iterations` is reached
- Return a tuple `(final_x, iterations_used)` where `final_x` is the final x value and `iterations_used` is how many iterations it took. (this will be equal to the number of times you called the logistic map function)

**Physics Note**: For r < 1, the system dies out. For 1 < r < 3, it converges to a fixed point. For r > 3, more complex behaviors emerge.

Return a tuple of (float, int).

```python
# Provided helper function for the logistic map
def logistic_map(x, r):
    """Apply one iteration of the logistic map: x_next = r * x * (1 - x)"""
    return r * x * (1 - x)

def find_fixed_point(initial_x, r, tolerance=0.001, max_iterations=1000):
    # Write your code here!
```

```python
grader.check("q4b")
```

## Question 5: Package Usage and Documentation

SciPy is one of the most important packages for scientific computing in Python. You'll use two key modules:
- `scipy.integrate` for numerical integration
- `scipy.special` for special mathematical functions

**Documentation resources:**
- [SciPy Integration Guide](https://docs.scipy.org/doc/scipy/tutorial/integrate.html)
- [SciPy Special Functions](https://docs.scipy.org/doc/scipy/reference/special.html)

Note: The datahub environment already has scipy in the environment, if you were doing this on a local environment, you would have to run something like $pip install scipy or $conda install scipy in the command line to be able to import scipy.

### Q5a: Numerical Integration (10 points)
Write a function `calculate_work_done(force_func, x_start, x_end)` that calculates the work done by a variable force over a displacement using numerical integration.

In physics, work done by a variable force F(x) from position x₁ to x₂ is:
$$W = \int_{x_1}^{x_2} F(x) \, dx$$

Your function should:
- Take a force function `force_func` (that takes position x and returns force)
- Take starting position `x_start` and ending position `x_end`
- Use `scipy.integrate.quad` to compute the integral
- Return the work done in Joules

**Hint:** You'll need to import `scipy.integrate.quad`. Look at the documentation to understand how to use it!

```python
# Write your code here! 
#import relevant package and define function
```

```python
grader.check("q5a")
```

### Q5b: Special Functions (10 points)

Write a function `quantum_harmonic_oscillator(n, x)` that calculates the normalized wavefunction of a quantum harmonic oscillator using Hermite polynomials from `scipy.special`.

The normalized wavefunction is:
$$\psi_n(x) = \left(\frac{m\omega}{\pi\hbar}\right)^{1/4} \frac{1}{\sqrt{2^n n!}} H_n\left(\sqrt{\frac{m\omega}{\hbar}}x\right) e^{-\frac{m\omega x^2}{2\hbar}}$$

For simplicity, we'll use dimensionless units where $\sqrt{\frac{m\omega}{\hbar}} = 1$, so:
$$\psi_n(x) = \frac{1}{\sqrt{2^n n! \sqrt{\pi}}} H_n(x) e^{-x^2/2}$$

Your function should:
- Take quantum number `n` (0, 1, 2, ...) and position `x`
- Use scipy to evaluate the physicist's Hermite polynomial (not probabilist's)
- Use scipy for the factorial
- Return the wavefunction value ψₙ(x)

```python
# Write your code here!
```

```python
grader.check("q5b")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> classify_projectile(0, 10) == 'vertical_up'
True
```
```
>>> classify_projectile(5, 0) == 'horizontal'
True
```
```
>>> classify_projectile(3, 4) == 'oblique_up'
True
```


Tests for q1b (points: None)
```
>>> range_val, flight_time = calculate_projectile_motion(10, 9.8)
>>> abs(range_val - 20.0) < 1e-10 and abs(flight_time - 2.0) < 1e-10
True
```
```
>>> calculate_projectile_motion(5, 0) == (0, 0)
True
```
```
>>> range_val, flight_time = calculate_projectile_motion(0, 19.6)
>>> range_val == 0 and abs(flight_time - 4.0) < 1e-10
True
```


Tests for q1c (points: None)
```
>>> assess_projectile_safety(10, 9.8, max_range=100, max_height=50) == 'SAFE'
True
```
```
>>> assess_projectile_safety(5, 0) == 'SAFE'
True
```
```
>>> assess_projectile_safety(0, 0) == 'SAFE'
True
```


Tests for q2a (points: None)
```
>>> result = analyze_oscillation(1.0, 4, 1)
>>> expected = [1.0, 0.0, -1.0, 0.0]
>>> all((abs(result[i] - expected[i]) < 1e-10 for i in range(4)))
True
```
```
>>> result = analyze_oscillation(2.0, 1, 1)
>>> abs(result[0] - 2.0) < 1e-10
True
```


Tests for q2b (points: None)
```
>>> result = analyze_wave_interference([1], [2], time=0.25)
>>> abs(result - 2.0) < 1e-10
True
```
```
>>> result = analyze_wave_interference([1, 2], [1, 2], time=0.125)
>>> abs(result - (1.0 * math.sin(math.pi / 4) + 2.0 * math.sin(math.pi / 2))) < 1e-10
True
```


Tests for q3a (points: None)
```
>>> fibonacci(0) == 0 and fibonacci(1) == 1
True
```
```
>>> fibonacci(2) == 1 and fibonacci(3) == 2 and (fibonacci(4) == 3)
True
```


Tests for q3b (points: None)
```
>>> recursive_sum([]) == 0
True
```
```
>>> recursive_sum([5]) == 5 and recursive_sum([1, 2, 3]) == 6 and (recursive_sum([0, 0, 0]) == 0)
True
```


Tests for q4a (points: None)
```
>>> radioactive_decay(1000, 5730, 100) == (62, 4)
True
```
```
>>> radioactive_decay(100, 1000, 100) == (100, 0)
True
```


Tests for q4b (points: None)
```
>>> result = find_fixed_point(0.5, 2.5, tolerance=0.001)
>>> abs(result[0] - 0.6) < 0.01 and result[1] < 100
True
```
```
>>> result = find_fixed_point(0.8, 0.5, tolerance=0.001)
>>> abs(result[0] - 0.0) < 0.01
True
```
```
>>> result = find_fixed_point(0.3, 4.0, tolerance=0.001, max_iterations=50)
>>> result[1] == 50
True
```


Tests for q5a (points: None)
```
>>> def test_spring_force(x):
...     return -10 * x
>>> spring_work = calculate_work_done(test_spring_force, 0, 0.5)
>>> abs(spring_work - -1.25) < 1e-10
True
```
```
>>> def test_constant_force(x):
...     return 5.0
>>> constant_work = calculate_work_done(test_constant_force, 2, 7)
>>> abs(constant_work - 25.0) < 1e-10
True
```


Tests for q5b (points: None)
```
>>> psi_0_0 = quantum_harmonic_oscillator(0, 0)
>>> expected_0_0 = math.pi ** (-0.25)
>>> bool(abs(psi_0_0 - expected_0_0) < 1e-10)
True
```
```
>>> psi_1_0 = quantum_harmonic_oscillator(1, 0)
>>> bool(abs(psi_1_0 - 0) < 1e-10)
True
```
