```python
# Initialize Otter
import otter
grader = otter.Notebook("week01_demo.ipynb")
```

# Week 01: Introduction to Python & Jupyter Notebook

Welcome to Physics 77/88! In this demo, we'll cover:

- Basic Python syntax and conventions
- How to use Jupyter notebooks for coding
- A mock homework problem with Otter-style public tests


## 1. Basic Python Syntax

- **Comments** start with `#` and are ignored by Python.
- **Indentation** is significant: use 4 spaces per indentation level.
- Statements end at the end of the line (no semicolons needed).


```python
# This is a comment
print('Hello, world!')
```

```python
x = 5  # integer
pi = 3.14159  # float
name = 'Alice'  # string
print(f'x = {x}, pi = {pi}, name = {name}')
```

## 2. Control Flow & Indentation

Use `if` statements and loops to control execution:

```python
if x > 0:
    print('x is positive')
else:
    print('x is zero or negative')
```

## 3. Jupyter Notebook Basics

- Run a cell: **Shift + Enter**
- Insert a cell: **A** (above) or **B** (below)
- Delete a cell: **DD**
- Change cell type: **Y** (code) or **M** (markdown)
- Common magic commands:
  - `%pwd`: show current working directory
  - `%ls`: list files in directory


```python
%pwd
```

## 4. Mock Homework Problem

**Question 1:** Write a function `square` that returns the square of its argument.

```python
def square(x):
    return ...
```

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

These are some submission instructions.

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> square(1) == 1
True
```
```
>>> square(0) == 0
True
```
