```python
# Initialize Otter
import otter
grader = otter.Notebook("ws0.ipynb")
```

## Notes:
- Please don't import external packages I haven't imported as part of the assignment, as this will cause otter grader to crash upon submission, resulting in a zero, this is fixable but not ideal.
- To turn in the assignment, run all code and test, save the notebook, then run the final cell to generate a .zip file to submit on gradescope.
- AI tools are allowed, but please disclose the way you used AI tools in the AI disclosure section (ungraded). 

# Week 01 Workshop: Introduction to Python & Jupyter Notebook

## **Exercise 1: Multiplier** 

At the end of the demo notebook, we saw a function that returned the square of a number. Here is another sample function that returns the sum of two numbers.

```python
def adder(a,b):
    return a + b
```

The function below takes two numbers. Complete the function so it returns the mutiple of the two.

```python
def multiplier(a,b):
    ...
    # Write your code here! Don't forget to return your answer.
```

```python
grader.check("q1")
```

## **Exercise 2:** Adding strings together

There are a few different ways to join strings in python. One way is to simply add them.

```python
string1 = "Beginning of the"
string2 = "sentence."

totalString = string1 + " " + string2
print(totalString)

#if you copy/paste this into a code cell, it will print: Beginning of the sentence.
```

We can also use the string.join() method. Instead of typing string, we replace that with the string that we want to use to separate the strings we are combining. We also need to have all of the strings in an iterable format. For now, we'll ignore what that means and just place them in square brackets. This makes them into a list.

```python
string1 = "first half +"
string2 = "second half"
iterableStrings = [string1,string2]
totalString = " ".join(iterableStrings)
print(totalString)

#if you copy/paste this into a code cell, it will print: first half + second half
#if you change the second to last line to totalString = "this is what goes in between".join(iterableStrings)
#it will print: first half +this is what goes in betweensecond half
```

Complete the function below using the method of your choice to add together two strings. The function also takes the character that should separate them. It should retrun one string that is the first string then the seperator, then the second string.

If the function call is `stringCombiner("abc", "123", " ")` it should return `"abc 123"`

```python
def stringCombiner(string1,string2,separator):
    return combinedString
```

```python
grader.check("q2")
```

## **Exercise 3:** All math results

Complete the function so it returns the two numbers added, subtracted, multiplied, and divided. Try to include a statement that checks if the second number is 0 to avoid division issues. If you realize you are dividing by 0, set divided to "Error: division by zero."

```python
def calculator(a,b):
    return added, subtracted, multiplied, divided
```

```python
grader.check("q3")
```

## **Exercise 4:** Mixing Strings and Math
Look up an easy way to repeat a string given a certain number of times. This can be done in one line! Name the function repeatString and pass it two positional arguments, the first one the string to be repeated, and the second one the number of times to repeat the string. 

```python
def repeatString(phrase,repeatTimes):
    repeatedString = ...
    return repeatedString
```

```python
grader.check("q4")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> multiplier(10, 3) == 30
True
```
```
>>> multiplier(0.5, 0.25) == 0.125
True
```
```
>>> multiplier(0, 0) == 0
True
```


Tests for q2 (points: None)
```
>>> stringCombiner('Hello', 'World', ' ') == 'Hello World'
True
```
```
>>> stringCombiner('Physics77', 'Workshop 1', '__ :) __') == 'Physics77__ :) __Workshop 1'
True
```
```
>>> stringCombiner('test', 'combiner', ' string ') == 'test string combiner'
True
```


Tests for q3 (points: None)
```
>>> calculator(1, 2) == (3, -1, 2, 0.5)
True
```
```
>>> calculator(4, 0) == (4, 4, 0, 'Error: division by zero.')
True
```


Tests for q4 (points: None)
```
>>> repeatString('Hello, World!', 3) == 'Hello, World!Hello, World!Hello, World!'
True
```
```
>>> repeatString('Test 2', 0) == ''
True
```
