```python
# Initialize Otter
import otter
grader = otter.Notebook("hw1.ipynb")
```

## Notes:
- Please don't import external packages I haven't imported as part of the assignment, as this will cause otter grader to crash upon submission, resulting in a zero, this is fixable but not ideal.
- To turn in the assignment, run all code and test, save the notebook, then run the final cell to generate a .zip file to submit on gradescope.
- AI tools are allowed, but please disclose the way you used AI tools in the AI disclosure section (ungraded). 

# Homework 1

```python
# Setup (standard library only)
# You may import math if needed in your solutions.
import math
```

## **Question 1**: Numbers, Units, and Floating-Point

### Q1a: Cyclotron frequency
Given charge q (C), magnetic field B (T), and electron mass m_e (kg), compute the cyclotron angular frequency ω_c = q B / m_e in rad/s.

Implement `electron_cyclotron_omega(q_C, B_T, m_e_kg)` and return a float.

```python
def electron_cyclotron_omega(q_C, B_T, m_e_kg):
    # BEGIN SOLUTION NO PROMPT
    return q_C * B_T / m_e_kg
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q1a")
```

### Q1b: Unit conversion (Tesla to Gauss)
1 tesla = 10,000 gauss. Implement `tesla_to_gauss(B_T)` that returns the equivalent value in gauss (float).

```python
def tesla_to_gauss(B_T):
    # BEGIN SOLUTION NO PROMPT
    return B_T * 10_000
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q1b")
```

### Q1c: Floating-point near-equality
Implement `nearly_equal(a, b, tol)` that returns True if |a - b| <= tol, otherwise False.

```python
def nearly_equal(a, b, tol):
    # BEGIN SOLUTION NO PROMPT
    return abs(a - b) <= tol
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q1c")
```

## **Question 2**: lists, tuples

### Q2a: Basic list stats
Implement `basic_stats(values)` returning a tuple `(mean, vmin, vmax)` using built-ins `sum`, `len`, `min`, `max`. If `values` is empty, return `(None, None, None)`.

```python
def basic_stats(values):
    # BEGIN SOLUTION NO PROMPT
    if len(values) == 0:
        return (None, None, None)
    mean = sum(values)/len(values)
    minimum = min(values)
    maximum = max(values)
    return (mean, minimum, maximum)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (mean, minimum, maximum)
    """; # END PROMPT
```

```python
grader.check("q2a")
```

### Q2b: Unique species
Implement `unique_species(observed)` where `observed` is a list like `["H", "He", "H", "O", "N", "N"]`. Return a tuple `(species_set, count)` where `species_set` is a set of unique items and `count` is its length. For the aforementioned
list, the output would be `(("H", "He", "N", "O"), 4)`

```python
def unique_species(observed):
    # BEGIN SOLUTION NO PROMPT
    s = set(observed)
    return (s, len(s))
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q2b")
```

## **Question 3**: Time-series slicing (lists)

Given a uniformly sampled list `samples` (e.g., acceleration readings), implement `slice_time_series(samples, first_n, step)` that returns a tuple of two lists:
1) the first `first_n` samples, and
2) every `step`-th sample starting at index 0 (i.e. if step=2, `[1,2,3,4,5,6,7,8,9,10]` would become `[1,3,5,7,9]`).

Use only list slicing/indexing; do not use loops.

```python
def slice_time_series(samples, first_n, step):
    # BEGIN SOLUTION NO PROMPT
    first = samples[:first_n]
    stepped = samples[::step]
    return (first, stepped)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (first, stepped)
    """; # END PROMPT
```

```python
grader.check("q3")
```

## **Question 4:** F-strings!

We want to calculate how far a ball will drop in Earth's gravity in a certain amount of time. To help with grading, assume g = 9.8 m/s and use SI units for your calculations. Complete the function below and have it return a string saying `"The ball will fall ___ m in ___ seconds."` Replace the blanks with your values, and be careful about spacing. Use an f-string! Don't round your answer. Look at how it compares to what you calculate.

Write a second function called `roundedDistance` that is similar to your first function, but takes an argument `digits` with a default value of 1. Use a built in function to round the value of distance that you return.

IMPORTANT: The rouynding here isn't with the builtin `round()` function. You need to implement the rounding in the f-string, and digits will refer to digits after the decimal in what is put into the string. 

```python
def distance(time):
    # BEGIN SOLUTION NO PROMPT
    g = 9.8 # m/s
    dist = 0.5 * g * time**2
    return f"The ball will fall {dist} m in {time} seconds."
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
def roundedDistance(time,digits = 1): # SOLUTION
    # BEGIN SOLUTION NO PROMPT
    g = 9.8 # m/s
    dist = 0.5 * g * time**2
    return f"The ball will fall {dist:.{digits}f} m in {time} seconds."
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q4")
```

## **Question 5:** Defaults in functions and special characters

In this exercise, you will write a function called `fancy_message` that returns the input message in quotes followed by a new line. If no message is put, set the default `message = "Hello World"`. The function should also allow the user to specify a number of times to repeat the message. If no number is specified, set `repeat = 1`. (Note: loops aren't required for this problem, look it up if you aren't sure how.)

As an example, if you have `message = "message to repeat"` and set `repeat = 2`, you should see:

```python
"message to repeat"
"message to repeat"
```

```python
# BEGIN SOLUTION NO PROMPT
def fancy_message(message="Hello, World!", repeat=1):
    str = f'\"{message}\"\n' * repeat
    return str
# END SOLUTION
""" # BEGIN PROMPT
# Write your code here!
"""; # END PROMPT
```

```python
grader.check("q5")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the raw cell below.

```python
# BEGIN SOLUTION NO PROMPT
# END SOLUTION
""" # BEGIN PROMPT
"""
# write ai disclosure here:

"""
"""; # END PROMPT
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> abs(electron_cyclotron_omega(1.602176634e-19, 0.75, 9.1093837015e-31) - 1.602176634e-19 * 0.75 / 9.1093837015e-31) < 1e-05
True
```
```
>>> abs(electron_cyclotron_omega(1.0, 2.0, 4.0) - 0.5) < 1e-05
True
```
```
>>> abs(electron_cyclotron_omega(0.0, 5.0, 1.0) - 0.0) < 1e-05
True
```


Tests for q1b (points: None)
```
>>> abs(tesla_to_gauss(0.0025) - 25.0) < 1e-05
True
```
```
>>> abs(tesla_to_gauss(0.0) - 0.0) < 1e-05
True
```
```
>>> abs(tesla_to_gauss(0.0001234) - 1.234) < 1e-05
True
```


Tests for q1c (points: None)
```
>>> nearly_equal(0.1 + 0.2, 0.3, 1e-12) == True
True
```
```
>>> nearly_equal(55, 50, 1) == False
True
```
```
>>> nearly_equal(5.0, 5.0 + 1e-06, 1e-07) == False
True
```


Tests for q2a (points: None)
```
>>> basic_stats([4.98, 5.02, 5.01, 4.99, 5.03]) == (5.006, 4.98, 5.03)
True
```
```
>>> basic_stats([]) == (None, None, None)
True
```
```
>>> basic_stats([1, 1, 1, 1, 1, 1]) == (1, 1, 1)
True
```
```
>>> basic_stats([1000000.0, -1000000.0, 0, 0, 0, 0]) == (0, -1000000.0, 1000000.0)
True
```


Tests for q2b (points: None)
```
>>> s, n = unique_species(['H', 'He', 'H', 'O'])
>>> s == {'H', 'He', 'O'} and n == 3
True
```
```
>>> unique_species([]) == (set(), 0)
True
```
```
>>> unique_species(['H', 'He', 'H', 'O', 'N', 'N']) == ({'H', 'He', 'N', 'O'}, 4)
True
```


Tests for q3 (points: None)
```
>>> samples = [0.0, 0.5, 1.0, 0.5, 0.0, -0.5, -1.0, -0.5]
>>> first, stepped = slice_time_series(samples, 3, 2)
>>> first == [0.0, 0.5, 1.0] and stepped == [0.0, 1.0, 0.0, -1.0]
True
```
```
>>> samples = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
>>> first, stepped = slice_time_series(samples, 6, 8)
>>> first == [0, 1, 2, 3, 4, 5] and stepped == [0, 8]
True
```
```
>>> samples = [0.0, 0.5, 1.0, 0.5, 0.0, -0.5, -1.0, -0.5]
>>> first, stepped = slice_time_series(samples, 20, 3)
>>> first == samples and stepped == [0.0, 0.5, -1.0]
True
```
```
>>> samples = [0.0, 0.5, 1.0, 0.5, 0.0, -0.5, -1.0, -0.5]
>>> first, stepped = slice_time_series(samples, 0, 1)
>>> first == [] and stepped == samples
True
```


Tests for q4 (points: None)
```
>>> distance(0) == 'The ball will fall 0.0 m in 0 seconds.'
True
```
```
>>> distance(2.5) == 'The ball will fall 30.625000000000004 m in 2.5 seconds.'
True
```
```
>>> roundedDistance(2.5) == 'The ball will fall 30.6 m in 2.5 seconds.'
True
```
```
>>> roundedDistance(10, 4) == 'The ball will fall 490.0000 m in 10 seconds.'
True
```


Tests for q5 (points: None)
```
>>> fancy_message() == '"Hello, World!"\n'
True
```
```
>>> fancy_message(message='Hello', repeat=2) == '"Hello"\n"Hello"\n'
True
```
```
>>> fancy_message(repeat=3) == '"Hello, World!"\n"Hello, World!"\n"Hello, World!"\n'
True
```
```
>>> fancy_message('test message') == '"test message"\n'
True
```
